Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python

```

---

# Método da Decomposição LU
Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (26 pontos)

Nessa parte você deve implementar os métodos das substituições sucessivas e o método da decomposição LU. Inicie pelo método das Substituições Sucessivas. 

<!-- grade_id: subst_suc, points: 0, grade: false, solution: true -->
```python
def substituicoes_sucessivas(A, b):
    '''Executa o método das substituições sucessivas para resolver o sistema 
       linear triangular inferior Ax=b.
       Parâmetros de entrada: A é uma matriz triangular inferior e b é o vetor constante. 
       Saída: vetor x
    '''
    # ESCREVA SEU CODIGO AQUI
    n = len(A)
    y = n*[0]
    for i in range(0,n,1):
        s = sum([A[i][j]*y[j] for j in range(0,i)])
        y[i] = (b[i]-s)/A[i][i]
    return y
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte resultado:
```
[2.0, -1.0, 5.0, 3.0]

```

```python
A = [[2, 0, 0, 0],
     [3, 5, 0, 0],
     [1, -6, 8, 0],
     [-1, 4, -3, 9]]
b = [4, 1, 48, 6]
x = substituicoes_sucessivas(A, b)
print(x)
```

Output:
```
[2.0, -1.0, 5.0, 3.0]
```

<!-- grade_id: subst_suc_test, points: 6, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função substituicoes_sucessivas retorna a saída correta para outras entradas. 
Outros testes podem ser realizados durante a correção."""
A = [[3, 0, 0, 0],
     [2, -3, 0, 0],
     [1, 0, 5, 0],
     [0, 2, 4, -3]]
b = [6, 7, -8, -3]
assert substituicoes_sucessivas(A,b) == [2.0, -1.0, -2.0, -2.3333333333333335]
A = [[-1, 0, 0],
     [2, 3, 0],
     [-1, 4, -5]]
b = [-1, 8, -8]
assert substituicoes_sucessivas(A,b) == [1.0, 2.0, 3.0]
```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de AssertionError indica que a função está implementada de modo incorreto. 

## Decomposição LU

Implemente uma função que decompõe uma matriz em A no produto de duas matrizes LU.

Para isso, use a sua função identidade criada no notebook trab-gauss_jordan.ipynb. Copie e cole a sua função na célula abaixo:

<!-- grade_id: identidade, points: 0, grade: false, solution: true -->
```python
def identidade(n):
    '''Cria uma matriz identidade de ordem n.
    Parâmetros de entrada:  n é a ordem da matriz.
    Saída: matriz identidade de ordem n
    '''
    # ESCREVA SEU CODIGO AQUI
    
    # Criando uma matriz com zeros
    matriz = []
    for i in range(n):
        linha = []
        for j in range(n):
            linha.append(0)
        matriz.append(linha)
    
    #Criando a matriz identidade
    for i in range(n):
        for j in range(n):
            if i == j:
                matriz[i][j] = 1
    
    return matriz
```

Agora implemente a função lu.

<!-- grade_id: lu, points: 0, grade: false, solution: true -->
```python
def lu(A):
    '''
    Decompõe a matriz A no produto de duas matrizes L e U. Onde L é uma matriz 
    triangular inferior unitária e U é uma matriz triangular superior.
    Parâmetros de entrada: A é uma matriz quadrada de ordem n.
    Saída: (L,U) tupla com as matrizes L e U
    '''
    # ESCREVA SEU CODIGO AQUI
    n = len(A)
    L = identidade(n)
    U = [linha[:] for linha in A]
    for j in range(n-1):
        for i in range(j+1,n):
            L[i][j] = U[i][j]/U[j][j]
            for k in range(j+1,n):
                U[i][k] = U[i][k] - L[i][j]*U[j][k] 
            U[i][j] = 0
    return (L, U)
            
```

Agora teste a sua função com o exemplo visto em sala: 

```python
def formata_matriz(M):
    m = len(M) # número de linhas
    n = len(M[0]) # número de colunas
    s = ""
    for i in range(m):
        for j in range(n):
           s += "%9.3f " % M[i][j]
        s += "\n"
    return s

A = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
(L, U) = lu(A)
print("L: \n%s" % formata_matriz(L))
print("U: \n%s" % formata_matriz(U))
```

Output:
```
L: 
    1.000     0.000     0.000 
   -2.000     1.000     0.000 
    4.000     3.000     1.000 

U: 
    1.000    -3.000     2.000 
    0.000     2.000     3.000 
    0.000     0.000   -12.000
```

Ao executar a célula acima, você deve ver o seguinte resultado:

```
L: 
    1.000     0.000     0.000 
   -2.000     1.000     0.000 
    4.000     3.000     1.000 

U: 
    1.000    -3.000     2.000 
    0.000     2.000     3.000 
    0.000     0.000   -12.000 
```

<!-- grade_id: lu_test, points: 10, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função lu retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
assert lu(A1) == ([[1, 0, 0], [-2.0, 1, 0], [4.0, 3.0, 1]], [[1, -3, 2], [0, 2.0, 3.0], [0, 0, -12.0]])
A1 = [[6, -2, 2, 4],
     [12, -8, 6, 10],
     [3, -13, 9, 3],
     [-6, 4, 1, -18]]
assert lu(A1) == ([[1, 0, 0, 0], [2.0, 1, 0, 0], [0.5, 3.0, 1, 0], [-1.0, -0.5, 2.0, 1]], [[6, -2, 2, 4], [0, -4.0, 2.0, 2.0], [0, 0, 2.0, -5.0], [0, 0, 0, -3.0]])

```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de AssertionError indica que a função está implementada de modo incorreto.

### Resolução de sistema linear usando decomposição LU

Na célula abaixo, copie e cole a função `substituicoes_retroativas` implementada no notebook `trab-gauss.ipynb`.

<!-- grade_id: subst_retro, points: 0, grade: false, solution: true -->
```python
def substituicoes_retroativas(A, b):
    '''Executa o método das substituições retroativas para resolver o sistema 
       linear triangular superior Ax=b.
       Parâmetros de entrada: A é uma matriz triangular superior e b é o vetor constante. 
    '''
    # ESCREVA SEU CODIGO AQUI
    n = len(A)
    y = n*[0]
    for i in range(n-1,-1,-1):
        s = sum([A[i][j]*y[j] for j in range(i+1,n)])
        y[i] = (b[i]-s)/A[i][i]
    return y
```

Agora implemente a função **resolve_lu(A,b)**.

<!-- grade_id: resolve_lu, points: 0, grade: false, solution: true -->
```python
def resolve_lu(A,b):
    '''
    Executa o método LU para resolver o sistema  linear Ax=b.
    Esse método inicialmente decompõe a matriz em L e U e depois resolve os 
    dois sistemas lineares triangulares.
    Parâmetros de entrada: A é uma matriz quadrada de ordem n e b é o vetor constante.
    Saída: vetor x solução do sistema.
    '''
    # ESCREVA SEU CODIGO AQUI
    (L, U) = lu(A)
    # Ly = b
    y = substituicoes_sucessivas(L, b)
    # Ux = y
    x = substituicoes_retroativas(U, y)
    return x
    
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar as células acima antes de executar a célula abaixo), você deve ver o seguinte resultado:

```[2.0, -1.0, 3.0]```

```python
A = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b = [11, -15, 29]
x = resolve_lu(A,b)
print(x)
```

Output:
```
[2.0, -1.0, 3.0]
```

<!-- grade_id: resolve_lu_test, points: 10, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função resolve_lu retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b1 = [11, -15, 29]
assert resolve_lu(A1, b1) == [2.0, -1.0, 3.0]

A1 = [[6, -2, 2, 4],
     [12, -8, 6, 10],
     [3, -13, 9, 3],
     [-6, 4, 1, -18]]
b1 = [16, 26, -19, -34]
assert resolve_lu(A1, b1) == [3.0, 1.0, -2.0, 1.0]

```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de **`AssertionError`** indica que a função está implementada de modo incorreto.

---
## Parte B (10 pontos)

![Figura 1: O problema dos reatores](https://emanueles.github.io/metodos/images/reatores.png)


A figura 1 mostra cinco reatores conectados por tubos. A razão de fluxo de massa através de cada tubo é calculada como o produto de fluxo (Q) por concentração (c). Num estado estável, o fluxo de massa entrando e saindo de um reator deve ser o mesmo. Por exemplo, para o primeiro reator o equilíbrio de massa pode ser escrito como $Q_{01}c_{01} + Q_{31}c_3 = Q_{15}c_1 + Q_{12}c_1$. Escreva as equações de equilíbrio de massa para os demais reatores e calcule as concentrações em cada reator utilizando a sua função `resolve_lu`.  

<!-- grade_id: reatores, points: 0, grade: false, solution: true -->
```python
## Defina a sua matriz A e o vetor b e chame a sua função resolve_lu. 
## Guarde o resultado na variável x
# ESCREVA SEU CODIGO AQUI
# Reator 1 - Q01c01 + Q31c3 = Q15c1 + Q12c1  // 9c1 - 3c3 = 120
# Reator 2 - Q12c1 = Q23c2 + Q24c2 + Q25c2   // 4c1 - 4c2 = 0
# Reator 3 - Q03c03 + Q23c2 = Q31c3 + Q34c3  // 2c2 - 9c3 = -350
# Reator 4 - Q24c2 + Q34c3 + Q54c5 = Q44c4   // c2 + 6c3 - 9c4 + 2c5 = 0
# Reator 5 - Q15c1 + Q25c2 = Q54c5 + Q55c5   // 5c1 + c2 - 6c5 = 0
A = [[9,0,-3,0,0],
    [4,-4,0,0,0],
    [0,2,-9,0,0],
    [0,1,6,-9,2],
    [5,1,0,0,-6]]
b = [120,0,-350,0,0]
x = resolve_lu(A,b)
print(x)
```

Output:
```
[28.4, 28.4, 45.199999999999996, 39.599999999999994, 28.400000000000002]
```

<!-- grade_id: reatores_test, points: 10, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula. 
## Certifique-se de que a célula acima foi executada logo antes de executar essa célula
```