Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python
NOME = "Matheus dos Santos Bastos"
EQUIPE = "Equipe 05"
```

---

# Método de Gauss-Jacobi
Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (22 pontos)

Nessa parte você deve implementar o método de Gauss-Jacobi. Antes disso, você precisa implementar a função `norma` para calcular a norma de vetores.

<!-- grade_id: norma, points: 0, grade: false, solution: true -->
```python
def norma(v,x):
    """Calcula a norma entre dois vetores v e x.
    """
    n = len(v)
    
    num = 0
    den = 0
    z = 0
    
    while z < n:
        difNum = abs(v[z] - x[z])
        
        if difNum > num:
            num = difNum
        if abs(v[z]) > den:
            den = abs(v[z])
            
        z += 1
    
    return num/den
```

Agora precisamos testar se a função está implementada corretamente. Iremos usar o exemplo mostrado em sala.

$x^{(0)} = \left(\begin{array}{r}
0.7\\
-1.6\\
0.6\\
\end{array}\right)\ e\ x^{(1)}= \left(\begin{array}{r}
0.96\\
-1.86\\
0.94\\
\end{array}\right) $

Ao executar a célula abaixo, você deverá ver o seguinte resultado:
```
0.18279569892473116
```

```python
x0 = [0.7, -1.6, 0.6]
x1 = [0.96, -1.86, 0.94]
d = norma(x1,x0)
print(d)
```

Output:
```
0.18279569892473116
```

<!-- grade_id: norma_test, points: 10, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função norma retorna a saída correta para outras entradas. 
Outros testes podem ser realizados durante a correção."""

x0 = [0, 0, 0, 0]
x1 = [2, 3, 4, 5]
assert norma(x1,x0) == 1.0

x0 = [0.7, -1.6, 0.6]
x1 = [0.96, -1.86, 0.94]
assert norma(x1,x0) == 0.18279569892473116

```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de **`AssertionError`** indica que a função está implementada de modo incorreto. 

Agora implemente o método da eliminação de Gauss-Jacobi.

<!-- grade_id: jacobi, points: 0, grade: false, solution: true -->
```python
def jacobi(A, b, epsilon, iterMax=50):
    """Resolve o sistema linear Ax=b usando o método iterativo Gauss-Jacobi.
    O critério de parada utiliza a norma-infinito.
    Saída é o vetor x.
    
    """
    n = len(A)
    x = n * [0]
    v = n * [0]
    k = 0
    l = 0
    
    while l < n:
        x[l] = b[l]/A[l][l]
        l += 1
    
    for c in range(0,n):
        for d in range(0,n):
            if c!= d:
                A[c][d] = A[c][d]/A[c][c]
        
        b[c] = b[c]/A[c][c]
        A[c][c] = 0 

    while k < iterMax:
        i = 0
        j = 0
    
        while i < n:
            AxSum = 0
            cont = 0
            
            while cont < n:
                if cont != i:
                    AxSum += -1 * A[i][cont] * x[cont]
                    
                cont += 1
        
            v[i] = (b[i] + AxSum)
            i += 1
        
        d = norma(v,x)
    
        if d < epsilon:
            return v
        
        while j < n:
            x[j] = v[j]
            j += 1
        
        k += 1
    print("Erro: Limite de iterações atingido!")
    return(x)
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte resultado:

```[0.9994, -1.9888000000000001, 0.9984]```

```python
A = [[10, 2 ,1],
     [1, 5, 1],
     [2, 3, 10]]
b = [7, -8, 6]
eps = 0.05
x = jacobi(A,b,eps)
print(x)
```

Output:
```
[0.9994, -1.9888000000000001, 0.9984]
```

<!-- grade_id: jacobi_test, points: 12, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função jacobi retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A1 = [[10, 2 ,1],
     [1, 5, 1],
     [2, 3, 10]]
b1 = [7, -8, 6]
eps = 0.00001
assert jacobi(A1, b1, eps) == [1.000001683824, -1.999997692128, 1.0000024307039999]
A1 = [[10, 3 ,-2],
     [2, 8, -1],
     [1, 1, 5]]
b1 = [57, 20, -4]
eps = 0.00001
assert jacobi(A1, b1, eps) == [5.0000014097125005, 1.00000782225, -1.9999973479000002]

# testando divergência
# jacobi deve retornar a última raiz calculada quando o método divergir
A1 = [[6, -2, 2, 4],
     [12, -8, 6, 10],
     [3, -13, 9, 3],
     [-6, 4, 1, -18]]
b1 = [16, 26, -19, -34]
eps = 0.001
assert jacobi(A1, b1, eps, 10) == [175.65036337449033, -366.97874928422397, 352.90185246106375, 69.3964852217211]
```

Output:
```
Erro: Limite de iterações atingido!
```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de **`AssertionError`** indica que a função está implementada de modo incorreto.

---
## Parte B (10 pontos)

![Figura 1: O problema dos reatores](https://emanueles.github.io/metodos/images/reatores.png)


A figura 1 mostra cinco reatores conectados por tubos. A razão de fluxo de massa através de cada tubo é calculada como o produto de fluxo (Q) por concentração (c). Num estado estável, o fluxo de massa entrando e saindo de um reator deve ser o mesmo. Por exemplo, para o primeiro reator o equilíbrio de massa pode ser escrito como $Q_{01}c_{01} + Q_{31}c_3 = Q_{15}c_1 + Q_{12}c_1$. Escreva as equações de equilíbrio de massa para os demais reatores e calcule as concentrações em cada reator utilizando a sua função `jacobi`, usando $\varepsilon = 0.000001$.  

<!-- grade_id: reatores, points: 0, grade: false, solution: true -->
```python
## Defina a sua matriz A e o vetor b e chame a sua função jacobi.
## Guarde o resultado na variável x
# ESCREVA SEU CODIGO AQUI

A = [[9, 0, -3, 0, 0],
     [4, -4, 0, 0, 0],
     [0, 2, -9, 0, 0],
     [0, 1, 6, -9, 2],
     [5, 1, 0, 0, -6]]

b = [120, 0, -350, 0, 0]

x = jacobi(A, b, 0.000001)
print(x)
```

Output:
```
[28.399995308456695, 28.399995308456695, 45.19999253317756, 39.59997831537849, 28.399990490263754]
```

<!-- grade_id: reatores_test, points: 10, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula. 
## Certifique-se de que a célula acima foi executada logo antes de executar essa célula
```