Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python

```

---

# Método de Gauss
Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (22 pontos)

Nessa parte você deve implementar os métodos das substituições retroativas, da Eliminação de Gauss e da Eliminação de Gauss modificado para retornar o determinante da matriz dos coeficientes. Inicie pelo método das Substituições Retroativas. 

<!-- grade_id: subst_retro, points: 0, grade: false, solution: true -->
```python
def substituicoes_retroativas(A, b):
    n = len(A)
    y = n*[0]
    for i in range(n-1,-1,-1):
        s = sum([A[i][j]*y[j] for j in range(i+1,n)])
        y[i] = (b[i] - s )/A[i][i]
    return y
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte resultado:
```
[-3.0, 0.0, 2.0, 4.0]
```

```python
A = [[5, -2, 6, 1],
      [0, 3, 7, -4],
      [0, 0, 4, 5],
      [0, 0, 0, 2]]
b = [1, -2, 28, 8]
x = substituicoes_retroativas(A,b)
print(x)
```

Output:
```
[-3.0, 0.0, 2.0, 4.0]
```

<!-- grade_id: subst_retro_test, points: 6, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função substituicoes_retroativas retorna a saída correta para outras entradas. 
Outros testes podem ser realizados durante a correção."""
A = [[5, -2, 6, 1],
      [0, 3, 7, -4],
      [0, 0, 4, 5],
      [0, 0, 0, 2]]
b = [1, -2, 28, 8]
assert substituicoes_retroativas(A,b) == [-3.0, 0.0, 2.0, 4.0]
A = [[2, 2, 3, 4],
     [0, 5, 6, 7],
     [0, 0, 8, 9],
     [0, 0, 0, 10]]
b = [20, 34, 25, 10]
assert substituicoes_retroativas(A,b) == [2.0, 3.0, 2.0, 1.0]
A = [[-1, 2, -1],
     [0, 3, 6],
     [0, 0, -5]]
b = [0, 24, -15]
assert substituicoes_retroativas(A,b) == [1.0, 2.0, 3.0]

```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de AssertionError indica que a função está implementada de modo incorreto. 

Agora implemente o método da eliminação de Gauss.

<!-- grade_id: gauss, points: 0, grade: false, solution: true -->
```python
def gauss(A, b):
    n = len(A)
    for i in range(0,n):
        for j in range(i+1,n):
            fatorm = A[j][i]/A[i][i]
            for aux in range(n):
                A[j][aux]=A[j][aux] - fatorm*A[i][aux]
            b[j] = b[j] - fatorm*b[i]
    x = substituicoes_retroativas(A,b)
    return (x)
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte resultado:

```[2.0, -1.0, 3.0]```

```python
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b1 = [11, -15, 29]
x = gauss(A1, b1)
print(x)
```

Output:
```
[2.0, -1.0, 3.0]
```

<!-- grade_id: gauss_test, points: 9, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função gauss retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b1 = [11, -15, 29]
assert gauss(A1, b1) == [2.0, -1.0, 3.0]
A1 = [[6, -2, 2, 4],
     [12, -8, 6, 10],
     [3, -13, 9, 3],
     [-6, 4, 1, -18]]
b1 = [16, 26, -19, -34]
assert gauss(A1, b1) == [3.0, 1.0, -2.0, 1.0]
```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de AssertionError indica que a função está implementada de modo incorreto.

Agora implemente o método da eliminação de Gauss modificado para retornar também o determinante da matriz dos coeficientes.

<!-- grade_id: gauss_det, points: 0, grade: false, solution: true -->
```python
def gauss_det(A, b):
    det = 1
    n = len(A)
    for i in range(0,n):
        det *= A[i][i]
        for j in range(i+1,n):
            fatorm = A[j][i]/A[i][i]
            for aux in range(n):
                A[j][aux]=A[j][aux] - fatorm*A[i][aux]
            b[j] = b[j] - fatorm*b[i]
    x = substituicoes_retroativas(A,b)
    return (x,det)
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte resultado:

```[2.0, -1.0, 3.0] -24.0```

```python
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b1 = [11, -15, 29]
(x, det) = gauss_det(A1, b1)
print(x, det)
```

Output:
```
[2.0, -1.0, 3.0] -24.0
```

<!-- grade_id: gauss_det_test, points: 7, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função gauss retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b1 = [11, -15, 29]
assert gauss_det(A1, b1) == ([2.0, -1.0, 3.0], -24.0)

A1 = [[6, -2, 2, 4],
     [12, -8, 6, 10],
     [3, -13, 9, 3],
     [-6, 4, 1, -18]]
b1 = [16, 26, -19, -34]
assert gauss_det(A1, b1) == ([3.0, 1.0, -2.0, 1.0], 144.0)

```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de AssertionError indica que a função está implementada de modo incorreto.

---
## Parte B (10 pontos)

![Figura 1: O problema dos reatores](https://emanueles.github.io/metodos/images/reatores.png)


A figura 1 mostra cinco reatores conectados por tubos. A razão de fluxo de massa através de cada tubo é calculada como o produto de fluxo (Q) por concentração (c). Num estado estável, o fluxo de massa entrando e saindo de um reator deve ser o mesmo. Por exemplo, para o primeiro reator o equilíbrio de massa pode ser escrito como $Q_{01}c_{01} + Q_{31}c_3 = Q_{15}c_1 + Q_{12}c_1$. Escreva as equações de equilíbrio de massa para os demais reatores e calcule as concentrações em cada reator utilizando a sua função `gauss`.  

<!-- grade_id: reatores, points: 0, grade: false, solution: true -->
```python
A = [[9 , 0 , -3 , 0 , 0],
     [-4 , 4 , 0 , 0 , 0],
     [0 , -2 , 9 , 0, 0],
     [0 , 1 , 6 , -9 , 2],
     [5 , 1 , 0 , 0 , -6]]
b = [120 , 0 , 350 , 0 , 0]
x = gauss(A,b)
print(x)
```

Output:
```
[28.4, 28.4, 45.199999999999996, 39.599999999999994, 28.400000000000002]
```

<!-- grade_id: reatores_test, points: 10, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula. 
## Certifique-se de que a célula acima foi executada logo antes de executar essa célula
```