Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python

```

---

# Método de Gauss com Pivotação Parcial
Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (26 pontos)

Nessa parte você deve implementar os métodos da Eliminação de Gauss com Pivotação Parcial e da Eliminação de Gauss com Pivotação modificado para retornar o determinante da matriz dos coeficientes. Na célula abaixo, copie e cole o método das substituições retroativas implementado no notebook trab-gauss.ipynb.

<!-- grade_id: subst_retro, points: 0, grade: false, solution: true -->
```python
def substituicoes_retroativas(A, b):
    n = len(A)
    y = n*[0]
    for i in range(n-1,-1,-1):
        s = sum([A[i][j]*y[j] for j in range(i+1,n)])
        y[i] = (b[i] - s )/A[i][i]
    return y
```

Agora implemente a função para escolher o pivô:

<!-- grade_id: escolhe_pivo, points: 0, grade: false, solution: true -->
```python
def escolhe_pivo(k, A, b):
    n = len(A)
    p = 0
    for i in range(k,n):
        if (abs(A[i][k]) > abs(A[k][k])):
            p += 1
            for j in range(n):                
                aux1 = A[k][j]
                A[k][j] = A[i][j]
                A[i][j] = aux1
            aux2 = b[k]
            b[k] = b[i]
            b[i] = aux2
    if p > 0:
        return True
    else :
        return False
```

Agora precisamos testar se a função está implementada corretamente. Iremos testar com o exemplo mostrado em sala. Se tudo estiver ok, você deverá ver o seguinte resultado:

```
True
A = [[3, 1, 0], [1, 2, 3], [0, 3, 4]]
b = [4, 3, 3]
```

```python
A = [[1, 2, 3],
     [3, 1, 0],
     [0, 3, 4]]
b = [3, 4, 3]
## Testaremos com a primeira coluna (k=0)
houveTroca = escolhe_pivo(0, A, b)
print(houveTroca)
print("A =", A)
print("b =", b)
```

Output:
```
True
A = [[3, 1, 0], [1, 2, 3], [0, 3, 4]]
b = [4, 3, 3]
```

Agora implemente a função `gauss_pivot` (ela deve chamar as funções `escolhe_pivo` e `substituicoes_retroativas`).

<!-- grade_id: gauss_pivot, points: 0, grade: false, solution: true -->
```python
def gauss_pivot(A, b):
    n = len(A)
    for i in range(0,n):
        escolhe_pivo(i,A,b)
        for j in range(i+1,n):
            fatorm = A[j][i]/A[i][i]
            for aux in range(n):
                A[j][aux]= A[j][aux] - fatorm*A[i][aux]
            b[j] = b[j] - fatorm*b[i]
    x = substituicoes_retroativas(A,b)
    return (x)
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar as células acima antes de executar a célula abaixo), você deve ver o seguinte resultado:
```
[1.0, 1.0, 0.0]
```

```python
A = [[1, 2, 3],
     [3, 1, 0],
     [0, 3, 4]]
b = [3, 4, 3]
x = gauss_pivot(A, b)
print(x)
```

Output:
```
[1.0, 1.0, 0.0]
```

<!-- grade_id: subst_retro_test, points: 13, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a gauss_pivot retorna a saída correta para outras entradas. 
Outros testes podem ser realizados durante a correção."""
A = [[0, -4, 5],
     [2, 3, 7],
     [5, -3, 2]]
b = [1, 12, 4]
assert gauss_pivot(A,b) == [1.0, 1.0, 1.0]
A = [[3, 2, 4],
    [1, 1, 2],
    [4, 3, -2]]
b = [1, 2, 3]
assert gauss_pivot(A,b) == [-3.0, 5.0, 0.0]
A = [[3, 2, 2, 4],
     [2, 3, 1, 2],
     [4, 1, 3, 5],
     [1, 4, 4, 3]]
b = [133, 90, 163, 126]
assert gauss_pivot(A,b) == [15, 8.0, 10.0, 13.0]
```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de AssertionError indica que a função está implementada de modo incorreto.

Agora implemente o método da eliminação de Gauss com pivotação modificado para retornar também o determinante da matriz dos coeficientes.

<!-- grade_id: gauss_det, points: 0, grade: false, solution: true -->
```python
def escolhe_pivo2(k, A, b):
    n = len(A)
    p = 0
    for i in range(k,n):
        if (abs(A[i][k]) > abs(A[k][k])):
            p += 1
            for j in range(n):                
                aux1 = A[k][j]
                A[k][j] = A[i][j]
                A[i][j] = aux1
            aux2 = b[k]
            b[k] = b[i]
            b[i] = aux2
    return p

def gauss_pivot_det(A, b):
    n = len(A)
    det = 1
    for i in range(0,n):
        p = escolhe_pivo2(i,A,b)
        det *= A[i][i] * (-1)**p
        for j in range(i+1,n):
            fatorm = A[j][i]/A[i][i]
            for aux in range(n):
                A[j][aux]= A[j][aux] - fatorm*A[i][aux]
            b[j] = b[j] - fatorm*b[i]
    x = substituicoes_retroativas(A,b)
    return (x,det)
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte reultado:

```[2.0, -1.0, 3.0] -24.0```

```python
A = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b = [11, -15, 29]
(x, det) = gauss_pivot_det(A, b)
print(x,det)
```

Output:
```
[2.0, -1.0, 3.0] -24.0
```

<!-- grade_id: gauss_det_test, points: 13, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função gauss_pivot_det retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b1 = [11, -15, 29]
assert gauss_pivot_det(A1, b1) == ([2.0, -1.0, 3.0], -24.0)

```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de **`AssertionError`** indica que a função está implementada de modo incorreto.

---
## Parte B (10 pontos)

![Figura 1: O problema dos reatores](https://emanueles.github.io/metodos/images/reatores.png)


A figura 1 mostra cinco reatores conectados por tubos. A razão de fluxo de massa através de cada tubo é calculada como o produto de fluxo (Q) por concentração (c). Num estado estável, o fluxo de massa entrando e saindo de um reator deve ser o mesmo. Por exemplo, para o primeiro reator o equilíbrio de massa pode ser escrito como $Q_{01}c_{01} + Q_{31}c_3 = Q_{15}c_1 + Q_{12}c_1$. Escreva as equações de equilíbrio de massa para os demais reatores e calcule as concentrações em cada reator utilizando a sua função `gauss_pivot`.  

<!-- grade_id: reatores, points: 0, grade: false, solution: true -->
```python
A = [[9 , 0 , -3 , 0 , 0],
     [-4 , 4 , 0 , 0 , 0],
     [0 , -2 , 9 , 0, 0],
     [0 , 1 , 6 , -9 , 2],
     [5 , 1 , 0 , 0 , -6]]
b = [120 , 0 , 350 , 0 , 0]
x = gauss_pivot(A,b)
print(x)
```

Output:
```
[28.4, 28.4, 45.199999999999996, 39.599999999999994, 28.400000000000002]
```

<!-- grade_id: reatores_test, points: 10, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula. 
## Certifique-se de que a célula acima foi executada logo antes de executar essa célula
```