Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python

```

---

# Método de Gauss-Jordan
Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (22 pontos)

Nessa parte você deve implementar o Método de Gauss-Jordan.

```python
def divide_pelo_pivo(k,A,b):
    n = len(A)
    for j in range (k+1, n):
        A[k][j] = A[k][j] / A[k][k]
    b[k] = b[k] / A[k][k]
    A[k][k] = 1
```

<!-- grade_id: gauss_jordan, points: 0, grade: false, solution: true -->
```python
def gauss_jordan(A, b):
    '''
    Executa o método de Gauss-Jordan para resolver o sistema linear Ax=b 
    transformando a matriz A na matriz identidade.
    Parâmetros de entrada: A é uma matriz quadrada de ordem n e b é o vetor constante.
    Saída: vetor solução x
    '''
    n = len(A)
    
    for k in range (0, n):
        divide_pelo_pivo(k, A, b)
    
        for i in range (0, n):
            if i != k:
                m = -A[i][k] / A[k][k]
                
                for j in range (k+1, n):
                    A[i][j] = A[i][j] + m*A[k][j]
                b[i] = b[i] + m*b[k]
                A[i][k] = 0
    return b
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte resultado:
```
[-2.9999999999999996, 5.0, 8.326672684688674e-17]
```

```python
A = [[3, 2, 4],
     [1, 1, 2],
     [4, 3, -2]]
b = [1, 2, 3]
x = gauss_jordan(A, b)
print(x)
```

Output:
```
[-2.9999999999999996, 5.0, 8.326672684688674e-17]
```

<!-- grade_id: gauss_jordan_test, points: 10, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função gauss_jordan retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A1 = [[1, -3, 2],
     [-2, 8, -1],
     [4, -6, 5]]
b1 = [11, -15, 29]
assert gauss_jordan(A1, b1) == [2.0, -1.0, 3.0]
A1 = [[6, -2, 2, 4],
     [12, -8, 6, 10],
     [3, -13, 9, 3],
     [-6, 4, 1, -18]]
b1 = [16, 26, -19, -34]
assert gauss_jordan(A1, b1) == [3.0, 1.0, -2.0, 1.0]
```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de AssertionError indica que a função está implementada de modo incorreto.

## Calculando a matriz inversa usando Gauss-Jordan

Antes de calcular a matriz inversa, implemente uma função para calcular a matriz identidade de ordem n.

<!-- grade_id: identidade, points: 0, grade: false, solution: true -->
```python
def identidade(n):
    '''Cria uma matriz identidade de ordem n.
    Parâmetros de entrada:  n é a ordem da matriz.
    Saída: matriz identidade de ordem n
    '''
    
    B = []
    
    for i in range (0,n):
        linha = []
        for j in range (0, n):
            if(i==j):
                linha.append(1)
            else:
                linha.append(0)
        B.append(linha)

    return B
```

Agora implemente a função que retorna a matriz inversa usando Gauss-Jordan:

```python
def divide_pelo_pivo_inv(k,A,Inv):
    n = len(A)

    for j in range (k+1, n):
        A[k][j] = A[k][j] / A[k][k]
    for j in range (0, n):
        Inv[k][j] = Inv[k][j] / A[k][k]
    A[k][k] = 1
```

<!-- grade_id: inversa, points: 0, grade: false, solution: true -->
```python
def gauss_inversa(A):
    '''
    Executa o método de Gauss-Jordan para calcular a matriz inversa de A.
    Parâmetros de entrada: A é uma matriz quadrada de ordem n.
    Saída: Inv é a matriz inversa de A.
    '''
    n = len(A)
    Inv = identidade(n)

    for k in range (0, n):
        divide_pelo_pivo_inv(k, A, Inv)
        
        for i in range (0, n):
            if i != k:
                m = -A[i][k]/A[k][k]

                for j in range (0, n):
                    A[i][j] = A[i][j] + m*A[k][j]
                    Inv[i][j] = Inv[i][j] + m*Inv[k][j]
                A[i][k] = 0
    return Inv
```

Agora teste a sua função com o exemplo visto em sala. Ao executar a célula abaixo (lembre-se de executar a célula acima antes de executar a célula abaixo), você deve ver o seguinte resultado:

```
Inv: 
1.5 1.5 -2.0 
-0.5 -1.5 1.0 
-0.5 -0.5 1.0 ```

```python
def formata_matriz(M):
    m = len(M) # número de linhas
    n = len(M[0]) # número de colunas
    s = ""
    for i in range(m):
        for j in range(n):
           s += "%s " % M[i][j]
        s += "\n"
    return s

A = [[2, 1, 3],
     [0, -1, 1],
     [1, 0, 3]]
Inv = gauss_inversa(A)
print("Inv: \n%s" % formata_matriz(Inv))
```

Output:
```
Inv: 
1.5 1.5 -2.0 
-0.5 -1.5 1.0 
-0.5 -0.5 1.0
```

<!-- grade_id: inversa_test, points: 12, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função gauss_inversa retorna a saída correta para outras entradas. Note que outros 
testes podem ser feitos durante a correção"""
A = [[2, -1, 3],
     [1, -2, 1],
     [0, 3, 2]]
assert gauss_inversa(A) == [[2.3333333333333335, -3.666666666666667, -1.6666666666666667], [0.6666666666666666, -1.3333333333333333, -0.3333333333333333], [-1.0, 2.0, 1.0]]

A1 = [[6, -2, 2, 4],
     [12, -8, 6, 10],
     [3, -13, 9, 3],
     [-6, 4, 1, -18]]
assert gauss_inversa(A1) == [[-3.4861111111111107, 2.1527777777777777, -0.6944444444444444, 0.3055555555555555], [8.291666666666666, -4.791666666666666, 1.4166666666666665, -0.5833333333333333], [11.916666666666666, -6.916666666666666, 2.1666666666666665, -0.8333333333333333], [3.6666666666666665, -2.1666666666666665, 0.6666666666666666, -0.3333333333333333]]

```

Se ao executar a célula acima não acontecer nada, os testes foram executados com sucesso. Qualquer mensagem de **`AssertionError`** indica que a função está implementada de modo incorreto.

---
## Parte B (10 pontos)

![Figura 1: O problema dos reatores](https://emanueles.github.io/metodos/images/reatores.png)


A figura 1 mostra cinco reatores conectados por tubos. A razão de fluxo de massa através de cada tubo é calculada como o produto de fluxo (Q) por concentração (c). Num estado estável, o fluxo de massa entrando e saindo de um reator deve ser o mesmo. Por exemplo, para o primeiro reator o equilíbrio de massa pode ser escrito como $Q_{01}c_{01} + Q_{31}c_3 = Q_{15}c_1 + Q_{12}c_1$. Escreva as equações de equilíbrio de massa para os demais reatores e calcule as concentrações em cada reator utilizando a sua função `gauss_jordan`.  

<!-- grade_id: reatores, points: 0, grade: false, solution: true -->
```python
## Defina a sua matriz A e o vetor b e chame a sua função gauss_jordan
## Guarde o resultado na variável x
# ESCREVA SEU CODIGO AQUI
C = [[9,0,-3,0,0],
     [4,-4,0,0,0],
     [0,2,-9,0,0],
     [0,1,6,-9,2],
     [5,1,0,0,-6]]

b = [120,0,-350,0,0]

x = gauss_jordan(C,b)
print(x)
```

Output:
```
[28.4, 28.4, 45.199999999999996, 39.599999999999994, 28.399999999999995]
```

<!-- grade_id: reatores_test, points: 10, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula. 
## Certifique-se de que a célula acima foi executada logo antes de executar essa célula
```