Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python

```

---

# Método da Bisseção

Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (15 pontos)

Implemente o método da Bisseção, tal que ele retorne uma tupla contendo um booleano e a raiz.

<!-- grade_id: bissecao, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, epsilon, maxIter):
    """Executa o método da bisseção para achar o zero de f no intervalo 
       [a,b] com precisão epsilon. O método executa no máximo maxIter
       iterações.
       Retorna uma tupla (houveErro, raiz), onde houveErro é booleano.
    """
    
    ##Inicializando f(a) e f(b)
    fa = f(a)
    fb = f(b)
    
    ##Checa se a função muda de sinal
    if fa * fb > 0:
        print("A função não muda de sinal.")
        return (True, None)
        

    ##Inicio da tabela com k, a, fa, b, fb, x, fx, IntervaloX
    print("k\t a\t\t fa\t\t b\t\t fb\t\t x\t\t fx\t\t intervalo")
        
    ##Inizializando o tamanho do intervalo e os valores de x e fx
    intervalo = abs(b-a)
    x = (a+b)/2
    fx = f(x)
        
    ##Dados de inicialização
    print("-\t%e\t%e\t%e\t%e\t%e\t%e\t%e" % (a, fa, b, fb, x, fx, intervalo))
       
    ##Checa se o intervalo é pequeno o bastante e retorna a raíz sem erros
    if intervalo <= epsilon:
        return (False, x)
        
    #Variável para auxiliar na iteração
    k = 1
        
    #Laço 
    while k <= maxIter:
        ##Checa se a raíz está entre a e x ou se está entre b e x e atualiza
        ##seus valores
        
        if fa * fx > 0:
            a = x
            fa = fx
        else:
            b = x
            fb = fx

        ##Atualiza o tamanho do intervalo e os novos valores de x e fx
        intervalo = abs(b-a)
        x = (a+b)/2
        fx = f(x)

        ##Dados da tabela atualizados
        print("%d\t%e\t%e\t%e\t%e\t%e\t%e\t%e" % (k, a, fa, b, fb, x, fx, intervalo))

        ##Checa se o novo intervalo é pequeno o suficiente e retorna a raíz, caso seja
        ##Teste de parada
        if intervalo <= epsilon:
            return (False, x)

        ##Incrementando o k
        k = k+1
        
    ##Mostra erro na tela se a função atingiu o número máximo de iterações e retorna 
    ##a última raíz
    print("Número máximo de iterações atingido!")
    return (True, x)
        
```

Teste a sua função com o exemplo visto em sala:
Compare o seu resultado com a tabela abaixo:
```
k	  a		  fa		  b		  fb		  x		  fx		intervX
-	0.000000e+00	3.000000e+00	1.000000e+00	-5.000000e+00	5.000000e-01	-1.375000e+00	1.000000e+00
1	0.000000e+00	3.000000e+00	5.000000e-01	-1.375000e+00	2.500000e-01	7.656250e-01	5.000000e-01
2	2.500000e-01	7.656250e-01	5.000000e-01	-1.375000e+00	3.750000e-01	-3.222656e-01	2.500000e-01
3	2.500000e-01	7.656250e-01	3.750000e-01	-3.222656e-01	3.125000e-01	2.180176e-01	1.250000e-01
4	3.125000e-01	2.180176e-01	3.750000e-01	-3.222656e-01	3.437500e-01	-5.313110e-02	6.250000e-02
5	3.125000e-01	2.180176e-01	3.437500e-01	-5.313110e-02	3.281250e-01	8.220291e-02	3.125000e-02
6	3.281250e-01	8.220291e-02	3.437500e-01	-5.313110e-02	3.359375e-01	1.447439e-02	1.562500e-02
7	3.359375e-01	1.447439e-02	3.437500e-01	-5.313110e-02	3.398438e-01	-1.934391e-02	7.812500e-03
8	3.359375e-01	1.447439e-02	3.398438e-01	-1.934391e-02	3.378906e-01	-2.438627e-03	3.906250e-03
9	3.359375e-01	1.447439e-02	3.378906e-01	-2.438627e-03	3.369141e-01	6.016918e-03	1.953125e-03
10	3.369141e-01	6.016918e-03	3.378906e-01	-2.438627e-03	3.374023e-01	1.788904e-03	9.765625e-04
```

```python
def f(x):
    return x**3 - 9*x + 3
a = 0
b = 1
epsilon = 0.001
maxIter = 20
(houveErro, raiz) = bissecao(f,a,b,epsilon,maxIter)
```

Output:
```
k	 a		 fa		 b		 fb		 x		 fx		 intervalo
-	0.000000e+00	3.000000e+00	1.000000e+00	-5.000000e+00	5.000000e-01	-1.375000e+00	1.000000e+00
1	0.000000e+00	3.000000e+00	5.000000e-01	-1.375000e+00	2.500000e-01	7.656250e-01	5.000000e-01
2	2.500000e-01	7.656250e-01	5.000000e-01	-1.375000e+00	3.750000e-01	-3.222656e-01	2.500000e-01
3	2.500000e-01	7.656250e-01	3.750000e-01	-3.222656e-01	3.125000e-01	2.180176e-01	1.250000e-01
4	3.125000e-01	2.180176e-01	3.750000e-01	-3.222656e-01	3.437500e-01	-5.313110e-02	6.250000e-02
5	3.125000e-01	2.180176e-01	3.437500e-01	-5.313110e-02	3.281250e-01	8.220291e-02	3.125000e-02
6	3.281250e-01	8.220291e-02	3.437500e-01	-5.313110e-02	3.359375e-01	1.447439e-02	1.562500e-02
7	3.359375e-01	1.447439e-02	3.437500e-01	-5.313110e-02	3.398438e-01	-1.934391e-02	7.812500e-03
8	3.359375e-01	1.447439e-02	3.398438e-01	-1.934391e-02	3.378906e-01	-2.438627e-03	3.906250e-03
9	3.359375e-01	1.447439e-02	3.378906e-01	-2.438627e-03	3.369141e-01	6.016918e-03	1.953125e-03
10	3.369141e-01
```

<!-- grade_id: correct_bissecao, points: 10, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função bissecao retorna a saída correta para várias entradas. Note que outros 
testes serão feitos durante a correção"""

assert bissecao(f,-4,-3,0.001,10) == (False,-3.15478515625)
assert bissecao(f,2,3,0.0001,14) == (False,2.816925048828125)

```

Output:
```
k	 a		 fa		 b		 fb		 x		 fx		 intervalo
-	-4.000000e+00	-2.500000e+01	-3.000000e+00	3.000000e+00	-3.500000e+00	-8.375000e+00	1.000000e+00
1	-3.500000e+00	-8.375000e+00	-3.000000e+00	3.000000e+00	-3.250000e+00	-2.078125e+00	5.000000e-01
2	-3.250000e+00	-2.078125e+00	-3.000000e+00	3.000000e+00	-3.125000e+00	6.074219e-01	2.500000e-01
3	-3.250000e+00	-2.078125e+00	-3.125000e+00	6.074219e-01	-3.187500e+00	-6.979980e-01	1.250000e-01
4	-3.187500e+00	-6.979980e-01	-3.125000e+00	6.074219e-01	-3.156250e+00	-3.604126e-02	6.250000e-02
5	-3.156250e+00	-3.604126e-02	-3.125000e+00	6.074219e-01	-3.140625e+00	2.879906e-01	3.125000e-02
6	-3.156250e+00	-3.604126e-02	-3.140625e+00	2.879906e-01	-3.148438e+00	1.265512e-01	1.562500e-02
7	-3.156250e+00	-3.604126e-02	-3.148438e+00	1.265512e-01	-3.152344e+00	4.539925e-02	7.812500e-03
8	-3.156250e+00	-3.604126e-02	-3.152344e+00	4.539925e-02	-3.154297e+00	4.715092e-03	3.906250e-03
9	-3.156250e+00	-3.604126e-02	-3.154297e+00	4.715092e-03	-3.155273e+00	-1.565406e-0
```

<!-- grade_id: bissecao_invalid_input, points: 5, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função bissecao retorna erro para entradas inválidas"""
# Erro! A função não muda de sinal:
assert bissecao(f,-2,-1,0.1,20) == (True, None)
assert bissecao(f,1,2,0.1,20) == (True, None)
# Erro! Número máximo de iterações atingido
assert bissecao(f,2,3,0.0000001,20) == (True,2.8169140815734863)

```

Output:
```
A função não muda de sinal.
A função não muda de sinal.
k	 a		 fa		 b		 fb		 x		 fx		 intervalo
-	2.000000e+00	-7.000000e+00	3.000000e+00	3.000000e+00	2.500000e+00	-3.875000e+00	1.000000e+00
1	2.500000e+00	-3.875000e+00	3.000000e+00	3.000000e+00	2.750000e+00	-9.531250e-01	5.000000e-01
2	2.750000e+00	-9.531250e-01	3.000000e+00	3.000000e+00	2.875000e+00	8.886719e-01	2.500000e-01
3	2.750000e+00	-9.531250e-01	2.875000e+00	8.886719e-01	2.812500e+00	-6.518555e-02	1.250000e-01
4	2.812500e+00	-6.518555e-02	2.875000e+00	8.886719e-01	2.843750e+00	4.034119e-01	6.250000e-02
5	2.812500e+00	-6.518555e-02	2.843750e+00	4.034119e-01	2.828125e+00	1.670418e-01	3.125000e-02
6	2.812500e+00	-6.518555e-02	2.828125e+00	1.670418e-01	2.820312e+00	5.041170e-02	1.562500e-02
7	2.812500e+00	-6.518555e-02	2.820312e+00	5.041170e-02	2.816406e+00	-7.515848e-03	7.812500e-03
8	2.816406e+00	-7.515848e-03	2.820312e+00	5.041170e-02	2.818359e+00	2.141567e-02	3.906250e-03
9	2.816406e+00	-7.515848e-03	2.818359e+00	2.141567e-02
```

---

## Parte B (5 pontos)

Usando a sua função `bissecao`, encontre a raiz de $f2(x) = 5x^7-4x^6+x^5-2x^4+3x^2-3x-4$. Usando uma técnica de isolamento de raízes, defina os valores de a e b, tal que a raiz seja encontrada com no máximo 10 iterações para epsilon = 0.0001.

<!-- grade_id: f2, points: 0, grade: false, solution: true -->
```python
def f2(x):
    return 5*x**7 - 4*x**6 + x**5 - 2*x**4 + 3*x**2 - 3*x - 4

```

f2(1) deve ser igual a -4. Tenha certeza de que implementou a função f2 corretamente. 

```python
f2(1)
```

Output:
```
-4
```

<!-- grade_id: correct_bissecao_f2, points: 0, grade: false, solution: true -->
```python
maxIter = 10
epsilon = 0.0001
## Defina a e b abaixo e depois chame a sua função bissecao para f2
## Mostre a raiz encontrada
## Utilizando a tabela como técnica de isolamento de raízes, nota-se que a função muda de sinal entre 1 e 2, pois
## testando a função com valores negativos o sinal continua negativo, mas para 2, a função fica poisitva. Mas, para
## ser mais preciso, podemos diminuir ainda mais o intervalo. 
## a = 1 => f(a) = -4 ; a = 1.1 => f(a) = -2.3303484999999973 ; b = 1.2 => f(b) = 1.0330880000000011
## A função já muda de sinal entre 1.1 e 1.2
a = 1.1
b = 1.2
(houveErro, raiz) = bissecao(f2, a, b, epsilon, maxIter)
if houveErro:
    print("O método retornou um erro.")
if raiz is not None:
    print("Raíz encontrada: %e" % raiz)
```

Output:
```
k	 a		 fa		 b		 fb		 x		 fx		 intervalo
-	1.100000e+00	-2.330348e+00	1.200000e+00	1.033088e+00	1.150000e+00	-9.212990e-01	1.000000e-01
1	1.150000e+00	-9.212990e-01	1.200000e+00	1.033088e+00	1.175000e+00	-2.134526e-02	5.000000e-02
2	1.175000e+00	-2.134526e-02	1.200000e+00	1.033088e+00	1.187500e+00	4.853260e-01	2.500000e-02
3	1.175000e+00	-2.134526e-02	1.187500e+00	4.853260e-01	1.181250e+00	2.270132e-01	1.250000e-02
4	1.175000e+00	-2.134526e-02	1.181250e+00	2.270132e-01	1.178125e+00	1.016091e-01	6.250000e-03
5	1.175000e+00	-2.134526e-02	1.178125e+00	1.016091e-01	1.176562e+00	3.982814e-02	3.125000e-03
6	1.175000e+00	-2.134526e-02	1.176562e+00	3.982814e-02	1.175781e+00	9.165793e-03	1.562500e-03
7	1.175000e+00	-2.134526e-02	1.175781e+00	9.165793e-03	1.175391e+00	-6.108608e-03	7.812500e-04
8	1.175391e+00	-6.108608e-03	1.175781e+00	9.165793e-03	1.175586e+00	1.523869e-03	3.906250e-04
9	1.175391e+00	-6.108608e-03	1.175586e+00	1.523869e-03	1.175488e+00	-2.293550e-03	1.953125e-04
10	1.175488e+00	
```

<!-- grade_id: test-bissecao-f2, points: 5, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula
```