Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python

```

---

# Método de Newton
Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (15 pontos)

Implemente o método de Newton, tal que ele retorne uma tupla contendo um booleano e a raiz.

<!-- grade_id: newton, points: 0, grade: false, solution: true -->
```python
def newton(f, flin, x0, epsilon, iterMax):
    """Executa o método de Newton-Raphson para achar o zero de f  
       a partir da derivada de f flin, aproximação inicial x0 
       e tolerância epsilon.
       Retorna uma tupla (houveErro, raiz), onde houveErro é booleano.
    """
    k = 0
    x1 = 0
    if abs(f(x0)) <= epsilon :
        return (False,x0)
    print("k\t x\t\t f(x)")
    print("%d\t%e\t%e" %(k,x0,f(x0)))
    while k != iterMax :
        if flin(x0) != 0 :
            x1 = x0 - f(x0)/flin(x0)  #Equação do método de Newton
            print("%d\t%e\t%e" %(k+1,x1,f(x1)))
            if abs(f(x1)) <= epsilon :
                return (False,x1)
            x0 = x1
            k += 1
            if abs(flin(x1)) <= epsilon :
                print("Erro! f'(x0) perto de 0, número máximo de iterações atingido")
                return (True,x1)
        else :
            print("Erro! f'(x0) = 0, divisão por zero")
            return (True,None)
    print("Erro! Número máximo de iterações atingido")
    return (True,x1)
```

Teste a sua função com o exemplo visto em sala: Compare o seu resultado com a tabela abaixo:
```
k	      x		     f(x)
0	5.000000e-01	-1.375000e+00
1	3.333333e-01	3.703704e-02
2	3.376068e-01	1.834089e-05
```

```python
def f(x):
    return x**3 - 9*x + 3
def flin(x):
    return 3*x**2 -9
x0 = 0.5
epsilon = 0.0001
maxIter = 20
(houveErro, raiz) = newton(f,flin, x0,epsilon,maxIter)
```

Output:
```
k	 x		 f(x)
0	5.000000e-01	-1.375000e+00
1	3.333333e-01	3.703704e-02
2	3.376068e-01	1.834089e-05
```

<!-- grade_id: correct_newton, points: 10, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função newton retorna a saída correta para várias entradas. Note que outros 
testes serão feitos durante a correção"""

assert newton(f, flin, -3.5, 0.0001,3) == (False,-3.154523328824589)
assert newton(f, flin, 2.5 ,0.00001,4) == (False,2.8169140527555467)

```

Output:
```
k	 x		 f(x)
0	-3.500000e+00	-8.375000e+00
1	-3.198198e+00	-9.288960e-01
2	-3.155363e+00	-1.752597e-02
3	-3.154523e+00	-6.675674e-06
k	 x		 f(x)
0	2.500000e+00	-3.875000e+00
1	2.897436e+00	1.247442e+00
2	2.820364e+00	5.117533e-02
3	2.816921e+00	1.002622e-04
4	2.816914e+00	3.875620e-10
```

<!-- grade_id: newton_invalid_input, points: 5, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função newton retorna erro para entradas inválidas"""
# Erro! f'(x0) = 0, divisão por zero
def ft(x):
    return x**2+x-6
def ftlin(x):
    return 2*x +1
assert newton(ft, ftlin, -1/2,0.1,20) == (True, None)
# Erro! f'(x0) perto de 0, número máximo de iterações atingido
assert newton(f, flin, 1.732,0.1,20) == (True, -6.711767508445506)
# Erro! Número máximo de iterações atingido
assert newton(f,flin, 2.5,0.0000001,3) == (True,2.8169208248794226)
```

Output:
```
k	 x		 f(x)
0	-5.000000e-01	-6.250000e+00
Erro! f'(x0) = 0, divisão por zero
k	 x		 f(x)
0	1.732000e+00	-7.392305e+00
1	-1.399885e+04	-2.743321e+12
2	-9.332564e+03	-8.128358e+11
3	-6.221709e+03	-2.408402e+11
4	-4.147807e+03	-7.136007e+10
5	-2.765205e+03	-2.114372e+10
6	-1.843471e+03	-6.264804e+09
7	-1.228982e+03	-1.856237e+09
8	-8.193226e+02	-5.499954e+08
9	-5.462175e+02	-1.629610e+08
10	-3.641487e+02	-4.828439e+07
11	-2.427713e+02	-1.430624e+07
12	-1.618558e+02	-4.238724e+06
13	-1.079163e+02	-1.255809e+06
14	-7.196279e+01	-3.720190e+05
15	-4.800320e+01	-1.101791e+05
16	-3.204428e+01	-3.261283e+04
17	-2.142643e+01	-9.640861e+03
18	-1.438043e+01	-2.841406e+03
19	-9.732989e+00	-8.314195e+02
20	-6.711768e+00	-2.389446e+02
Erro! Número máximo de iterações atingido
k	 x		 f(x)
0	2.500000e+00	-3.875000e+00
1	2.897436e+00	1.247442e+00
2	2.820364e+00	5.117533e-02
3	2.816921e+00	1.002622e-04
Erro! Número máximo de iterações atingido
```

---
## Parte B (5 pontos)

Usando a sua função `newton`, encontre a raiz de $f2(x) = 5x^7-4x^6+x^5-2x^4+3x^2-3x-4$. Usando uma técnica de isolamento de raízes, defina o valor de x0, tal que a raiz seja encontrada com no **máximo 3 iterações** para epsilon = 0.0001.

<!-- grade_id: f2, points: 0, grade: false, solution: true -->
```python
def f2(x):
    return 5*x**7 - 4*x**6 + x**5 -2*x**4 + 3*x**2 - 3*x - 4
def f2lin(x):
    return 35*x**6 - 24*x**5 + 5*x**4 + 6*x - 3
```

f2(1) deve ser igual a -4. Tenha certeza de que implementou a função f2 corretamente. 

```python
f2(1)
```

Output:
```
-4
```

<!-- grade_id: newton_f2, points: 0, grade: false, solution: true -->
```python
maxIter = 3
epsilon = 0.0001
## Defina x0 abaixo e depois chame a sua função newton para f2
## Mostre a raiz encontrada
x0 = 1.2
print(newton(f2,f2lin, x0,epsilon,maxIter))

```

Output:
```
k	 x		 f(x)
0	1.200000e+00	1.033088e+00
1	1.182596e+00	2.817741e-01
2	1.177387e+00	7.236191e-02
3	1.176012e+00	1.819596e-02
Erro! Número máximo de iterações atingido
(True, 1.1760117323035426)
```

<!-- grade_id: test-newton_f2, points: 5, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula
```