Antes de submeter esse arquivo, certifique-se de que tudo está funcionando conforme esperado. Para isso, **reinicie o Kernel** (na barra de menu, selecione Kernel$\rightarrow$Restart) e então **execute todas as células** (na barra de menu, selecione Cell$\rightarrow$Run All).

Tenha certeza de substituir todos os lugares onde está escrito `SEU CODIGO AQUI` ou `SUA RESPOSTA AQUI` pelo seu código ou sua resposta, bem como preencher com o seu nome e a sua equipe nas variáveis abaixo:

```python

```

---

# Método de Newton para Polinômios
Para essa parte do trabalho, usaremos o seguinte notebook jupyter:

---
## Parte A (15 pontos)

Implemente o método de Newton para polinômios, tal que ele retorne uma tupla contendo um booleano e a raiz.

<!-- grade_id: newton, points: 0, grade: false, solution: true -->
```python
def newton_poli(a, x, epsilon, iterMax=50):
    """Executa o método de Newton-Raphson para polinômios para achar o zero  
       do polinômio definido como coeficientes a, onde a = [a0, a1, ..., an].
       São necessárias a aproximação inicial x e a tolerância epsilon. 
       O método executa no máximo iterMax iterações.
       Retorna uma tupla (houveErro, raiz), onde houveErro é booleano.
    """
    n = len(a)-1
    print("k\t x\t\t p(x)")
    k = 0
    
    while k <= iterMax:    
        b = [1.0] * len(a)
        c = [1.0] * n

        i = n - 1
        while i != 0:
            b[i] = a[i] + b[i+1]*x
            c[i-1] = b[i] + c[i]*x
            i -= 1

        b[0] = a[0] + b[1]*x
   
        print("%d\t %e\t %e\t" %(k, x, b[0]))

        if abs(b[0]) <= epsilon:
            return(False, x)
        if abs(c[0]) == 0 :
            print("Erro! c = 0, divisão por zero")
            return (True, None)
        
        x = x - b[0]/c[0]
        k += 1

    print("Erro! Número máximo de iterações atingido")
    return(True, x)
```

Teste a sua função com o exemplo visto em sala: Compare o seu resultado com a tabela abaixo:
```
k	      x		     p(x)
0	5.000000e-01	-1.375000e+00
1	3.333333e-01	3.703704e-02
2	3.376068e-01	1.834089e-05
```

```python
a = [3, -9, 0, 1]
x0 = 0.5
epsilon = 0.0001
maxIter = 20
(houveErro, raiz) = newton_poli(a, x0, epsilon, maxIter)
```

Output:
```
k	 x		 p(x)
0	 5.000000e-01	 -1.375000e+00	
1	 3.333333e-01	 3.703704e-02	
2	 3.376068e-01	 1.834089e-05
```

<!-- grade_id: correct_newton, points: 10, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função newton retorna a saída correta para várias entradas. Note que outros 
testes serão feitos durante a correção"""

assert newton_poli(a, -3.5, 0.0001,3) == (False, -3.154523328824589)
assert newton_poli(a, 2.5 ,0.00001,4) == (False, 2.8169140527555467)

```

Output:
```
k	 x		 p(x)
0	 -3.500000e+00	 -8.375000e+00	
1	 -3.198198e+00	 -9.288960e-01	
2	 -3.155363e+00	 -1.752597e-02	
3	 -3.154523e+00	 -6.675674e-06	
k	 x		 p(x)
0	 2.500000e+00	 -3.875000e+00	
1	 2.897436e+00	 1.247442e+00	
2	 2.820364e+00	 5.117533e-02	
3	 2.816921e+00	 1.002622e-04	
4	 2.816914e+00	 3.875633e-10
```

<!-- grade_id: newton_invalid_input, points: 5, grade: true, solution: false -->
```python
# Não modifique essa célula
"""Verificando que a função newton retorna erro para entradas inválidas"""
# Erro! c = 0, divisão por zero
at = [-6, 1, 1]
assert newton_poli(at, -1/2,0.1,20) == (True, None)
# Erro! c perto de 0, número máximo de iterações atingido
assert newton_poli(a, 1.732,0.1,20) == (True, -4.817538586845097)
# Erro! Número máximo de iterações atingido
assert newton_poli(a, 2.5,0.0000001,3) == (True, 2.8169140527555467)
```

Output:
```
k	 x		 p(x)
0	 -5.000000e-01	 -6.250000e+00	
Erro! c = 0, divisão por zero
k	 x		 p(x)
0	 1.732000e+00	 -7.392305e+00	
1	 -1.399885e+04	 -2.743321e+12	
2	 -9.332564e+03	 -8.128358e+11	
3	 -6.221709e+03	 -2.408402e+11	
4	 -4.147807e+03	 -7.136007e+10	
5	 -2.765205e+03	 -2.114372e+10	
6	 -1.843471e+03	 -6.264804e+09	
7	 -1.228982e+03	 -1.856237e+09	
8	 -8.193226e+02	 -5.499954e+08	
9	 -5.462175e+02	 -1.629610e+08	
10	 -3.641487e+02	 -4.828439e+07	
11	 -2.427713e+02	 -1.430624e+07	
12	 -1.618558e+02	 -4.238724e+06	
13	 -1.079163e+02	 -1.255809e+06	
14	 -7.196279e+01	 -3.720190e+05	
15	 -4.800320e+01	 -1.101791e+05	
16	 -3.204428e+01	 -3.261283e+04	
17	 -2.142643e+01	 -9.640861e+03	
18	 -1.438043e+01	 -2.841406e+03	
19	 -9.732989e+00	 -8.314195e+02	
20	 -6.711768e+00	 -2.389446e+02	
Erro! Número máximo de iterações atingido
k	 x		 p(x)
0	 2.500000e+00	 -3.875000e+00	
1	 2.897436e+00	 1.247442e+00	
2	 2.820364e+00	 5.117533e-02	
3	 2.816921e+00	 1.002622e-04	
Erro! Número máximo de iteraçõe
```

---
## Parte B (5 pontos)

Usando a sua função `newton`, encontre a raiz de $f2(x) = 5x^7-4x^6+x^5-2x^4+3x^2-3x-4$. Usando uma técnica de isolamento de raízes, defina o valor de x0, tal que a raiz seja encontrada com no **máximo 3 iterações** para epsilon = 0.0001.

<!-- grade_id: newton_f2, points: 0, grade: false, solution: true -->
```python
maxIter = 3
epsilon = 0.0001
## Defina x0 e os coeficientes de f2 como a variável a abaixo e depois chame a sua função newton_poli
## Mostre a raiz encontrada
x0 = 1.2
a = [-4, -3, 3, -2, 1, -4, 5]
print(newton_poli(a, x0, epsilon, maxIter))
```

Output:
```
k	 x		 p(x)
0	 1.200000e+00	 -1.162970e+01	
1	 7.168402e-01	 -5.703046e+00	
2	 -5.635818e-01	 -6.380094e-01	
3	 -6.197341e-01	 5.728900e-02	
Erro! Número máximo de iterações atingido
(True, -0.6154821884518248)
```

<!-- grade_id: test-newton_f2, points: 5, grade: true, solution: false -->
```python
## Verificação da sua resposta. Não modifique essa célula
```